# Data Analysis Homework 3
### Due date: Friday, November 1st 2024, 11AM

```python
from __future__ import division
from IPython.display import HTML
from IPython.display import display
from scipy.optimize import *
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
```

## Question 1: Linear Regression, Curvature Matrix

Consider the data listed below,
\begin{equation}
\begin{array}{lcccccc}
\hline
{\rm frequency~(Hz)} &10&20&30&40&50&60\\
{\rm voltage~(mV)} &16&45&64&75&70&115\\
{\rm error~(mV)}   &5&5&5&5&30&5\\
\hline
{\rm frequency~(Hz)} &70&80&90&100&110&\\
{\rm voltage~(mV)} &142&167&183&160&221&\\
{\rm error~(mV)}   &5&5&5&30&5&\\
\hline
\end{array} 
\end{equation}

This data is also contained in the file 'linear_regression.csv'.

Required:
<br>
(i) Calculate the 4 elements of the curvature matrix.
<br>
(ii) Invert this to give the error matrix.
<br>
(iii) What are the uncertainties in the slope and intercept?
<br>
(iv) Comment on your answer.

### (i) Calculate the 4 elements of the curvature matrix.

<!-- grade_id: cell-a8061423011a8a5e, points: 0, grade: false, solution: true -->
```python
data = pd.read_csv('linear_regression.csv')
frequencies = data.iloc[:,0]
voltages = data.iloc[:,1]
voltage_errors = data.iloc[:,2]

def one_i():
    '''Your function should return something of the form np.matrix([[a_cc,a_cm],[a_mc,a_mm]])'''
    '''where m is the slope and c the intercept'''
    # YOUR CODE HERE
    return(curvature_matrix)

print(one_i())
```

<!-- grade_id: cell-205a22506ae95a24, points: 2, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
assert isinstance(one_i(), (list, tuple, np.ndarray)), \
    'Please make sure that the return value is list/array of floats.'
assert len(one_i()) == 2 , 'Please make sure you return a matrix with dimensions 2x2' 
```

### (ii) Invert this to give the error matrix.

<!-- grade_id: cell-f1181bd93715c586, points: 0, grade: false, solution: true -->
```python
data = pd.read_csv('linear_regression.csv')
frequencies = data.iloc[:,0]
voltages = data.iloc[:,1]
voltage_errors = data.iloc[:,2]

def one_ii():
    '''Your function should return something of the form np.matrix([[a_cc,a_cm],[a_mc,a_mm]])'''
    # YOUR CODE HERE
    return(inverted_matrix)

print(one_ii())
```

<!-- grade_id: cell-7ab85907fde2a15a, points: 2, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
assert isinstance(one_ii(), (list, tuple, np.ndarray)), \
    'Please make sure that the return value is list/array of floats.'
assert len(one_ii()) == 2 , 'Please make sure you return a matrix with dimensions 2x2' 
```

### (iii) What are the uncertainties in the slope and intercept?

<!-- grade_id: cell-6ff81eb1afbe53e7, points: 0, grade: false, solution: true -->
```python
data = pd.read_csv('linear_regression.csv')
frequencies = data.iloc[:,0]
voltages = data.iloc[:,1]
voltage_errors = data.iloc[:,2]

def one_iii():
    slope_uncertainty = 0
    intercept_uncertainty = 0
    # YOUR CODE HERE
    return(slope_uncertainty,intercept_uncertainty)

print(one_iii())
```

<!-- grade_id: cell-ea08060575a403ed, points: 2, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
assert isinstance(one_iii(), (list, tuple, np.ndarray)), \
    'Please make sure that the return value is list/array of floats.'
assert len(one_iii()) == 2 , 'Please make sure you return a list/array of length 2' 
```

<!-- grade_id: cell-961a84be93aa364a, points: 2, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
```

### (iv) Comment on your answer

<!-- grade_id: cell-f68bfa4e8b879207, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2: Using a calibration curve

A frequently encountered case where the correlation of the uncertainties must be taken into account is that of a calibration curve.  Consider the following set of measurements from an optical-activity experiment, where the angle of rotation of a plane-polarized light beam, $\theta$, is measured as a function of the independent variable, the concentration, $C$, of a sucrose solution. 

\begin{equation}
\begin{array}{lcccc}
\hline
C \mbox{ (g cm$^{-3}$)} &0.025&0.05&0.075&0.100\\
\theta \mbox{ (degrees)}&10.7&21.6&32.4&43.1\\
\hline
C \mbox{ (g cm$^{-3}$)}&0.125&0.150&0.175\\
\theta \mbox{ (degrees)}&53.9&64.9&75.4\\
\hline
\end{array} 
\end{equation}

The errors in the angle measurement are all $0.1^{\circ}$, the errors in the concentration are negligible.  A straight line  fit to the data yields  a gradient of $431.7\,^{\circ}\mbox{ g$^{-1}$ cm$^{3}$}$, and intercept $-0.03^{\circ}$. This data is contained in 'optical_activity.csv'.

<br>
Required:
<br>
(i) Show that the curvature matrix, $\mathsf{A}$, is given by 

\begin{equation}
\mathsf{A}=\left[\begin{array}{cc}
700\left((^{\circ})^{-2}\right)&70\left((^{\circ})^{-2}\mbox{g cm$^{-3}$}\right)\\
70\left((^{\circ})^{-2}\mbox{g cm$^{-3}$}\right)&8.75\left((\mbox{g/$^\circ$ cm$^{3})^2$}\right)\\
\end{array}\right] ,
\end{equation}


and that the error matrix  is 

\begin{equation}
\mathsf{C}=\left[\begin{array}{cc}
0.00714\left((^{\circ})^2\right)&-0.0571\left((^{\circ})^2\mbox{g$^{-1}$cm$^{3}$}\right)\\
-0.0571\left((^{\circ})^2\mbox{g$^{-1}$cm$^{3}$}\right)&0.571\left((^{\circ})^2\mbox{g$^{-2}$ cm$^{6}$}\right)\\
\end{array}\right] .
\end{equation}

The entry for the intercept is in the top left-hand corner, that for the gradient in the bottom right-hand corner.  
<br>
(ii) Calculate the associated correlation matrix.  

Use the  entries of the error matrix to answer the following  questions: 
<br>
(iii) What are the uncertainties in the best-fit intercept and gradient? 
<br>
(iv) What optical rotation is expected for a known concentration of $C=0.080g cm^{-3}$, and what is the uncertainty? 
<br>
(v) What is the concentration given a measured rotation of $\theta=70.3^{\circ}$ and what is the uncertainty?

### (i) Verify the curvature matrix and the error matrix above.

<!-- grade_id: cell-90c242e605b9f033, points: 0, grade: false, solution: true -->
```python
data = pd.read_csv('optical_activity.csv')
concentrations = data.iloc[:,0]
angles = data.iloc[:,1]
angle_errors = data.iloc[:,2]

def two_i():
    '''Your function should return something of the form np.matrix([[a_cc,a_cm],[a_mc,a_mm]]). Must return the curvature and error matricies)'''
    curvature_matrix = 0
    error_matrix = 0
    # YOUR CODE HERE
    return(curvature_matrix,error_matrix)

print(two_i())
```

<!-- grade_id: cell-e6bcb90497cb94c6, points: 2, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
assert isinstance(two_i(), (list, tuple, np.ndarray)), \
    'Please make sure that the return value is list/array of floats.'
assert len(two_i()) == 2 , 'Please make sure you return a list/array of two matrices' 
assert len(two_i()[0]) == 2, 'Please make sure that the first entry is a 2x2 matrix' 
assert len(two_i()[1]) == 2, 'Please make sure that the second entry is a 2x2 matrix'  
```

<!-- grade_id: cell-55f3298589f14581, points: 2, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
```

### (ii) Calculate the associated correlation matrix.  

<!-- grade_id: cell-6a7fa0e40768028d, points: 0, grade: false, solution: true -->
```python
data = pd.read_csv('optical_activity.csv')
concentrations = data.iloc[:,0]
angles = data.iloc[:,1]
angle_errors = data.iloc[:,2]

def two_ii():
    '''Your function should return something of the form np.matrix([[a_cc,a_cm],[a_mc,a_mm]])'''
    # YOUR CODE HERE
    
print(two_ii())    
```

<!-- grade_id: cell-c2b19f704f54c27f, points: 2, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
assert isinstance(two_ii(), (list, tuple, np.ndarray)), \
    'Please make sure that the return value is a matrix.'
assert len(two_ii()) == 2 
```

### (iii) What are the uncertainties in the best-fit intercept and gradient? 

<!-- grade_id: cell-f32635e579d41db7, points: 0, grade: false, solution: true -->
```python
data = pd.read_csv('optical_activity.csv')
concentrations = data.iloc[:,0]
angles = data.iloc[:,1]
angle_errors = data.iloc[:,2]

def two_iii():
    '''Your function should return the uncertainty in the gradient and intercept'''
    gradient_uncertainty = 0
    intercept_uncertainty = 0
    # YOUR CODE HERE
    return(gradient_uncertainty,intercept_uncertainty)

print(two_iii())
```

<!-- grade_id: cell-0e1f3d9729e50809, points: 1, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
assert isinstance(two_iii(), (list, tuple, np.ndarray)), \
    'Please make sure that the return value is a list/array of floats.'
assert len(two_iii()) == 2 
```

<!-- grade_id: cell-b94fe8908649d12a, points: 1, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
```

### (iv) What optical rotation is expected for a known concentration of $C=0.080g cm^{-3}$, and what is the uncertainty? 

<!-- grade_id: cell-48ce90c9edb9a1af, points: 0, grade: false, solution: true -->
```python
data = pd.read_csv('optical_activity.csv')
concentrations = data.iloc[:,0]
angles = data.iloc[:,1]
angle_errors = data.iloc[:,2]

def two_iv():
    '''Your function should return the angle and the uncertainty'''
    angle = 0
    uncertainty = 0
    # YOUR CODE HERE
    return(angle,uncertainty)

print(two_iv())
```

<!-- grade_id: cell-1d9ea692b906e00c, points: 1, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
assert isinstance(two_iv(), (list, tuple, np.ndarray)), \
    'Please make sure that the return value is a list/array of floats.'
assert len(two_iv()) == 2 
```

<!-- grade_id: cell-d8a2703a940a7583, points: 1, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
```

### (v) What is the concentration given a measured rotation of $\theta=70.3^{\circ}$ and what is the uncertainty? You must return your answer in $gcm^{-3}$

<!-- grade_id: cell-2a210b206abe39ca, points: 0, grade: false, solution: true -->
```python
data = pd.read_csv('optical_activity.csv')
concentrations = data.iloc[:,0]
angles = data.iloc[:,1]
angle_errors = data.iloc[:,2]

def two_v():
    '''Your function should return the concentration and uncertainty'''
    # YOUR CODE HERE
    return(concentration,uncertainty)

print(two_v())
```

<!-- grade_id: cell-fc6a0b34e288a268, points: 1, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
assert isinstance(two_v(), (list, tuple, np.ndarray)), \
    'Please make sure that the return value is a list/array of floats.'
assert len(two_v()) == 2 
```

<!-- grade_id: cell-173ca414fbffc47c, points: 1, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
```

## Question 3: Error bars from a $\chi^2$ minimisation to a non-linear function

In this question we will analyse the data shown in the figure below, which is an X-ray spectrum as a function of angle.

![title](diffraction.JPG)
 
The data is contained in the file 'LorentzianData.csv'. There are three columns: the angle, the signal (in counts per second), and the error.  The number of X-rays counted in 20 seconds was recorded.

The model to describe the data has four parameters:  the height of the Lorentzian lineshape, $S_0$; the angle at which the peak is centered, $\theta_{0}$;
 the angular width of the peak, $\Delta\theta$; and a constant background offset, $S_{\rm bgd}$. Mathematically, the signal, $S$, is of the form:
\begin{equation}
S=S_{\rm bgd}+\frac{S_{0}}{1+4\left(\frac{\theta-\theta_{0}}{\Delta\theta}\right)^2}.
\end{equation}

and the function is defined by lorentzian(theta, s_0, s_bgd,delta_theta,theta_0).

Required:
<br>
(i) Explain how the error in the count rate was calculated.
<br>
(ii) Perform a $\chi^2$ minimisation.  What are the best-fit parameters?
<br>
(iii) Evaluate the error matrix.
<br>
(iv) Calculate the correlation matrix.
<br>
(v) What are the uncertainties in the best-fit parameters?
<br>
(vi) If you can plot contour plots, show the $\chi^2$ contours for 
<br>
>(a) background vs. peak centre. 
<br>
>(b) background vs. peak width.  


These figures are shown in figure 6.11 of Hughes and Hase. Comment on the shape of the contours. Only your comment will be graded. 

### (i) Explain how the error in the count rate was calculated.

<!-- grade_id: cell-af660dec82aee39a, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### (ii) Perform a $\chi^2$ minimisation.  What are the best-fit parameters?

<!-- grade_id: cell-b803d7b2c879b923, points: 0, grade: false, solution: true -->
```python
data = pd.read_csv('LorentzianData.csv') 

def lorentzian(theta, s_0, s_bgd,delta_theta,theta_0):
    return s_bgd+(s_0/(1+4*(((theta-theta_0)/delta_theta)**2)))

def three_ii():
    s_0 = 0
    s_bgd = 0
    delta_theta = 0
    theta_0 = 0
    covariance_matrix = 0
    angles = data.iloc[:,0]
    intensity = data.iloc[:,1]
    intensity_errors = data.iloc[:,2]
    # YOUR CODE HERE
    return(s_0,s_bgd,delta_theta,theta_0,covariance_matrix)

print(three_ii())
```

<!-- grade_id: cell-e0ea799fe64f9f1a, points: 2, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
assert isinstance(three_ii(), (list, tuple, np.ndarray)), \
    'Please make sure that the return value is a list/array.'
assert len(three_ii()) == 5 , 'Please make sure that you return five values' 
```

### (iii) Evaluate the error matrix.

<!-- grade_id: cell-4190be1416227f9e, points: 0, grade: false, solution: true -->
```python
data = pd.read_csv('LorentzianData.csv') 

def three_iii():
    '''Your function should return something of the form np.matrix([4x4])'''
    # YOUR CODE HERE
    return(error_matrix)

print(three_iii())
```

<!-- grade_id: cell-1d1f5acb338ca208, points: 2, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
assert isinstance(three_iii(), (list, tuple, np.ndarray)), \
    'Please make sure that the return value is a matrix.'
assert len(three_iii()) == 4 , 'Please make sure that you return a 4x4 matrix' 
```

### (iv) Calculate the correlation matrix.

<!-- grade_id: cell-c7257914f67b1821, points: 0, grade: false, solution: true -->
```python
data = pd.read_csv('LorentzianData.csv') 

def three_iv():
    '''Your function should return something of the form np.matrix([[a_cc,a_cm],[a_mc,a_mm]])'''
    # YOUR CODE HERE
    return(correlation_matrix)

print(three_iv())
```

<!-- grade_id: cell-2fdddf67da129349, points: 2, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
assert isinstance(three_iv(), (list, tuple, np.ndarray)), \
    'Please make sure that the return value is a matrix.'
assert len(three_iv()) == 4 , 'Please make sure that you return a 4x4 matrix' 
```

### (v) What are the uncertainties in the best-fit parameters?

<!-- grade_id: cell-e6553187fe7ba54a, points: 0, grade: false, solution: true -->
```python
data = pd.read_csv('LorentzianData.csv') 

def three_v():
    uncertainty_s_0 = 0
    uncertainty_s_bgd = 0
    uncertainty_delta_theta = 0
    uncertainty_theta_0 = 0
    # YOUR CODE HERE
    return(uncertainty_s_0,uncertainty_s_bgd,uncertainty_delta_theta,uncertainty_theta_0)

print(three_v())
```

<!-- grade_id: cell-b82624e7d32c9049, points: 2, grade: true, solution: false -->
```python
'''TEST CELL - DO NOT DELETE'''
```

### (vi) These contours are shown in figure 6.11 of Hughes and Hase. Comment on the shape of the contours.

<!-- grade_id: cell-9d235bfae599fba1, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 4: Prove the following properties:

Assume in this question that the uncertainties in $A$ and $B$ are correlated.
<br>
(i) If $Z=A\pm B$, show that
${\displaystyle\alpha_{Z}^2=\alpha_{A}^2+\alpha_{B}^2\pm2\alpha_{AB}}$.
<br>
(ii) If $Z=A\times B$, show that
 ${\displaystyle\left(\frac{\alpha_Z}{Z}\right)^2=\left(\frac{\alpha_A}{A}\right)^2+\left(\frac{\alpha_B}{B}\right)^2+2\left(\frac{\alpha_{AB}}{AB}\right)}$.
<br>
(iii) If ${\displaystyle Z=\frac{A}{B}}$, show that
${\displaystyle\left(\frac{\alpha_Z}{Z}\right)^2=\left(\frac{\alpha_A}{A}\right)^2+\left(\frac{\alpha_B}{B}\right)^2-2\left(\frac{\alpha_{AB}}{AB}\right)}$.

<!-- grade_id: cell-98dd0155ce1eafc8, points: 6, grade: true, solution: true -->
YOUR ANSWER HERE